# Interact Exercise 6

## Imports

Put the standard imports for Matplotlib, Numpy and the IPython widgets in the following cell.

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import math
```

```python
from IPython.display import Image
from IPython.html.widgets import interact, interactive, fixed
```

## Exploring the Fermi distribution

In quantum statistics, the [Fermi-Dirac](http://en.wikipedia.org/wiki/Fermi%E2%80%93Dirac_statistics) distribution is related to the probability that a particle will be in a quantum state with energy $\epsilon$. The equation for the distribution $F(\epsilon)$ is:

```python
Image('fermidist.png')
```

In this equation:

* $\epsilon$ is the single particle energy.
* $\mu$ is the chemical potential, which is related to the total number of particles.
* $k$ is the Boltzmann constant.
* $T$ is the temperature in Kelvin.

In the cell below, typeset this equation using LaTeX:

<!-- grade_id: interactex06a, points: 2, grade: true, solution: true -->
\begin{equation*}
F(\epsilon)=\frac{1}{e^{(\epsilon-\mu)/kT}+1}
\end{equation*}

Define a function `fermidist(energy, mu, kT)` that computes the distribution function for a given value of `energy`, chemical potential `mu` and temperature `kT`. Note here, `kT` is a single variable with units of energy. Make sure your function works with an array and don't use any `for` or `while` loops in your code.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def fermidist(energy, mu, kT):
    """Compute the Fermi distribution at energy, mu and kT."""
    f=1/(math.exp((energy-mu)/kT)+1)
    return f
```

<!-- grade_id: interactex06b, points: 2, grade: true, solution: false -->
```python
assert np.allclose(fermidist(0.5, 1.0, 10.0), 0.51249739648421033)
assert np.allclose(fermidist(np.linspace(0.0,1.0,10), 1.0, 10.0),
    np.array([ 0.52497919,  0.5222076 ,  0.51943465,  0.5166605 ,  0.51388532,
               0.51110928,  0.50833256,  0.50555533,  0.50277775,  0.5       ]))
```

Output:
```
<ipython-input-23-6de4cc798fa3> in <module>()
      1 assert np.allclose(fermidist(0.5, 1.0, 10.0), 0.51249739648421033)
----> 2 assert np.allclose(fermidist(np.linspace(0.0,1.0,10), 1.0, 10.0),
      3     np.array([ 0.52497919,  0.5222076 ,  0.51943465,  0.5166605 ,  0.51388532,
      4                0.51110928,  0.50833256,  0.50555533,  0.50277775,  0.5       ]))

<ipython-input-22-936293db3af6> in fermidist(energy, mu, kT)
      1 def fermidist(energy, mu, kT):
      2     """Compute the Fermi distribution at energy, mu and kT."""
----> 3     f=1/(math.exp((energy-mu)/kT)+1)
      4     f=np.array(f)
      5     return f

TypeError: only length-1 arrays can be converted to Python scalars
```

Write a function `plot_fermidist(mu, kT)` that plots the Fermi distribution $F(\epsilon)$ as a function of $\epsilon$ as a line plot for the parameters `mu` and `kT`.

* Use enegies over the range $[0,10.0]$ and a suitable number of points.
* Choose an appropriate x and y limit for your visualization.
* Label your x and y axis and the overall visualization.
* Customize your plot in 3 other ways to make it effective and beautiful.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def plot_fermidist(mu, kT):
    plt.plot(fermidist,energy)
    plt.xlim(0, 10.0)
    plt.ylim(0, 10.0)
    plt.xlabel('Energy')
    plt.ylabel('Fermi Distribution')
    plt.tick_params(axis='y', direction='inout', length=10)
    plt.grid(True)
    plt.box(False)
```

```python
plot_fermidist(4.0, 1.0)
```

<!-- grade_id: interactex06c, points: 4, grade: true, solution: false -->
```python
assert True # leave this for grading the plot_fermidist function
```

Use `interact` with `plot_fermidist` to explore the distribution:

* For `mu` use a floating point slider over the range $[0.0,5.0]$.
* for `kT` use a floating point slider over the range $[0.1,10.0]$.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
interact(plot_fermidist, mu=widgets.FloatSlider(min=0,max=5,step=0.1,value=2.5), kT=widgets.FloatSlider(min=0.1,max=10.0,step=0.1,value=5.0));

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-25-a2c4edb35268> in <module>()
----> 1 interact(plot_fermidist, mu=widgets.FloatSlider(min=0,max=5,step=0.1,value=2.5), kT=widgets.FloatSlider(min=0.1,max=10.0,step=0.1,value=5.0));

NameError: name 'plot_fermidist' is not defined
```

Provide complete sentence answers to the following questions in the cell below:

* What happens when the temperature $kT$ is low?
* What happens when the temperature $kT$ is high?
* What is the effect of changing the chemical potential $\mu$?
* The number of particles in the system are related to the area under this curve. How does the chemical potential affect the number of particles.

Use LaTeX to typeset any mathematical symbols in your answer.

<!-- grade_id: interactex06d, points: 2, grade: true, solution: true -->
when temperature is low, F(E) goes to zero
when temperature is high, F(E) goes to infinity
