# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    x=list(x)
    y=list(y)
    count=[i for i in x if i in y]
    if len(x)!=len(y):
        return False
    elif len(count)==len(x):
        return True
    else:
        return False
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integer(x):
    a=str(2*int(x))
    b=str(3*int(x))
    c=str(4*int(x))
    d=str(5*int(x))
    e=str(6*int(x))
    if same_digits(x,a)==True and same_digits(a,b)==True and same_digits(b,c)==True and same_digits(c,d)==True and same_digits(d,e):
        return x
    else:
        return False
```

```python
for x in range(200000):
    smallest=0
    s=str(x)
    if integer(s)!=False:
        smallest+=x
        print(smallest)
    else:
        pass
```

Output:
```
0
142857
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```