# Algorithms Exercise 2

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import seaborn as sns
import numpy as np
```

## Peak finding

Write a function `find_peaks` that finds and returns the indices of the local maxima in a sequence. Your function should:

* Properly handle local maxima at the endpoints of the input array.
* Return a Numpy array of integer indices.
* Handle any Python iterable as input.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def find_peaks(a):
    """Find the indices of the local maxima in a sequence."""
    localmax=[]
    for x in range(len(a)):
        if x==0:
            if a[x]>a[1]:
                localmax.append(x)
            else:
                pass
        elif x==len(a):
            if a[x]>a[x-1]:
                localmax.append(x)
            else:
                pass
        elif a[x]>a[x-1] and a[x]>a[x+1]:
                localmax.append(x)
        else:
            pass
    return localmax
```

```python
find_peaks([2,0,1,0,2,0,1])
```

Output:
```
<ipython-input-37-e5e01639155c> in <module>()
----> 1 find_peaks([2,0,1,0,2,0,1])

<ipython-input-36-c0e8c7bcb471> in find_peaks(a)
     13             else:
     14                 pass
---> 15         elif a[x]>a[x-1] and a[x]>a[x+1]:
     16                 localmax.append(x)
     17         else:

IndexError: list index out of range
```

<!-- grade_id: algorithmsex02a, points: 5, grade: true, solution: false -->
```python
p1 = find_peaks([2,0,1,0,2,0,1])
assert np.allclose(p1, np.array([0,2,4,6]))
p2 = find_peaks(np.array([0,1,2,3]))
assert np.allclose(p2, np.array([3]))
p3 = find_peaks([3,2,1,0])
assert np.allclose(p3, np.array([0]))
```

Here is a string with the first 10000 digits of $\pi$ (after the decimal). Write code to perform the following:

* Convert that string to a Numpy array of integers.
* Find the indices of the local maxima in the digits of $\pi$.
* Use `np.diff` to find the distances between consequtive local maxima.
* Visualize that distribution using an appropriately customized histogram.

```python
from sympy import pi, N
pi_digits_str = str(N(pi, 10001))[2:]
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: algorithmsex02b, points: 5, grade: true, solution: false -->
```python
assert True # use this for grading the pi digits histogram
```