# Numpy Exercise 2

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Factorial

Write a function that computes the factorial of small numbers using `np.arange` and `np.cumprod`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def np_fact(n):
    """Compute n! = n*(n-1)*...*1 using Numpy."""
    if n == 0:
        return 1
    else:
        a = np.arange(1,n+1,1)
        b = a.cumprod(0)
        return b[n-1]
```

<!-- grade_id: numpyex02a, points: 4, grade: true, solution: false -->
```python
assert np_fact(0)==1
assert np_fact(1)==1
assert np_fact(10)==3628800
assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Write a function that computes the factorial of small numbers using a Python loop.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def loop_fact(n):
    """Compute n! using a Python for loop."""
    if n == 0:
        return 1
    else:
        factorial = 1
        for i in range(1,n+1):
            factorial *= i
        return factorial
```

<!-- grade_id: numpyex02b, points: 4, grade: true, solution: false -->
```python
assert loop_fact(0)==1
assert loop_fact(1)==1
assert loop_fact(10)==3628800
assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Use the `%timeit` magic to time both versions of this function for an argument of `50`. The syntax for `%timeit` is:

```python
%timeit -n1 -r1 function_to_time()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
%timeit -n1 -r1 np_fact(100)
```

Output:
```
1 loops, best of 1: 61 µs per loop
```

```python
%timeit -n1 -r1 loop_fact(100)
```

Output:
```
1 loops, best of 1: 26 µs per loop
```

In the cell below, summarize your timing tests. Which version is faster? Why do you think that version is faster?

<!-- grade_id: numpyex02c, points: 2, grade: true, solution: true -->
I would have guessed that np_fact was going to be faster, but it turned out loop_fact was faster.  I think np_fact was slower because it creates a whole array and then finds the product of all of the values in that array. loop_fact is faster because it iterates over a range and multiplies those numbers together as one step.

```python

```